Self-Other Generalisation Shapes Social Interaction and Is Disrupted in Borderline Personality Disorder

Author response:

The following is the authors’ response to the previous reviews.

Reviewer 1:

The authors frequently refer to their predictions and theory as being causal, both in the manuscript and in their response to reviewers. However, causal inference requires careful experimental design, not just statistical prediction. For example, the claim that "algorithmic differences between those with BPD and matched healthy controls" are "causal" in my opinion is not warranted by the data, as the study does not employ experimental manipulations or interventions which might predictably affect parameter values. Even if model parameters can be seen as valid proxies to latent mechanisms, this does not automatically mean that such mechanisms cause the clinical distinction between BPD and CON, they could plausibly also refer to the effects of therapy or medication. I recommend that such causal language, also implicit to expressions like "parameter influences on explicit intentional attributions", is toned down throughout the manuscript.

Thankyou for this chance to be clearer in the language. Our models and paradigm introduce a from of temporal causality, given that latent parameter distributions are directly influenced by latent parameter estimates at a previous point in time (self-uncertainty and other uncertainty directly governs social contagion). Nevertheless, we appreciate the reviewers perspective and have now toned down the language to reflect this.

Abstract:

‘Our model makes clear predictions about the mechanisms of social information generalisation concerning both joint and individual reward.’

Discussion:

‘We can simulate this by modelling a framework that incorporates priors based on both self and a strong memory impression of a notional other (Figure S3).’

‘We note a strength of this work is the use of model comparison to understand algorithmic differences between those with BPD and matched healthy controls.’

Although the authors have now much clearer outlined the stuy's aims, there still is a lack of clarity with respect to the authors' specific hypotheses. I understand that their primary predictions about disruptions to self-other generalisation processes underlying BPD are embedded in the four main models that are tested, but it is still unclear what specific hypotheses the authors had about group differences with respect to the tested models. I recommend the authors specify this in the introduction rather than refering to prior work where the same hypotheses may have been mentioned.

Thankyou for this further critique which has enabled us to more cleary refine our introduction. We have now edited our introduction to be more direct about our hypotheses, that these hypotheses are instantiated into formal models, and what our predictions were. We have also included a small section on how previous predictions from other computational assessments of BPD link to our exploratory work, and highlighted this throughout the manuscript.

‘This paper seeks to address this gap by testing explicitly how disruptions in self-other generalization processes may underpin interpersonal disruptions observed in BPD. Specifically, our hypotheses were: (i) healthy controls will demonstrate evidence for both self-insertion and social contagion, integrating self and other information during interpersonal learning; and (ii) individuals with BPD will exhibit diminished self-other integration, reflected in stronger evidence for observations that assume distinct self-other representations.

We tested these hypotheses by designing a dynamic, sequential, three-phase Social Value Orientation (Murphy & Ackerman, 2014) paradigm—the Intentions Game—that would provide behavioural signatures assessing whether BPD differed from healthy controls in these generalization processes (Figure 1A). We coupled this paradigm with a lattice of models (M1-M4) that distinguish between self-insertion and social contagion (Figure 1B), and performed model comparison:

M1. Both self-to-other (self-insertion) and other-to-self (social contagion) occur before and after learning M2. Self-to-other transfer only occurs M3. Other-to-self transfer only occurs M4. Neither transfer process, suggesting distinct self-other representations

We additionally ran exploratory analysis of parameter differences and model predictions between groups following from prior work demonstrating changes in prosociality (Hula et al., 2018), social concern (Henco et al., 2020), belief stability (Story et al., 2024a), and belief updating (Story, 2024b) in BPD to understand whether discrepancies in self-other generalisation influences observational learning. By clearly articulating our hypotheses, we aim to clarify the theoretical contribution of our findings to existing literature on social learning, BPD, and computational psychiatry.’

Caveats should also be added about the exploratory nature of the many parameter group comparisons. If there are any predictions about group differences that can be made based on prior literature, the authors should make such links clear.

Thank you for this. We have now included caveats in the text to highlight the exploratory nature of these group comparisons, and added direct links to relevant literature where able:

Introduction

‘We additionally ran exploratory analysis of parameter differences and model predictions between groups following from prior work demonstrating changes in prosociality (Hula et al., 2018), social concern (Henco et al., 2020), belief stability (Story et al., 2024a), and belief updating (Story, 2024b) in BPD to understand whether discrepancies in self-other generalisation influences observational learning. By clearly articulating our hypotheses, we aim to clarify the theoretical contribution of our findings to existing literature on social learning, BPD, and computational psychiatry.’

Model Comparison

‘We found that CON participants were best fit at the group level by M1 (Frequency = 0.59, Exceedance Probability = 0.98), whereas BPD participants were best fit by M4 (Frequency = 0.54, Exceedance Probability = 0.86; Figure 2A). This suggests CON participants are best fit by a model that fully integrates self and other when learning, whereas those with BPD are best explained as holding disintegrated and separate representations of self and other that do not transfer information back and forth.

We first explore parameters between separate fits (see Methods). Later, in order to assuage concerns about drawing inferences from different models, we examined the relationships between the relevant parameters when we forced all participants to be fit to each of the models (in a hierarchical manner, separated by group). In sum, our model comparison is supported by convergence in parameter values when comparisons are meaningful (see Supplementary Materials). We refer to both types of analysis below.’

Phase 2 analysis

‘Prior work predicts those with BPD should focus more intently on public social information, rather than private information that only concerns one party (Henco et al., 2020). In BPD participants, only new beliefs about the relative reward preferences – mutual outcomes for both player - of partners differed (see Fig 2E): new median priors were larger than median preferences in phase 1 (mean = -0.47; = -6.10, 95%HDI: -7.60, -4.60).’

‘Models of moral preference learning (Story et al., 2024) predicts that BPD vs non-BPD participants have more rigid beliefs about their partners. We found that BPD participants were equally flexible around their prior beliefs about a partner’s relative reward preferences (= -1.60, 95%HDI: -3.42, 0.23), and were less flexible around their beliefs about a partner’s absolute reward preferences (=-4.09, 95%HDI: -5.37, -2.80), versus CON (Figure 2B).’

Phase 3 analysis

‘Prior work predicts that human economic preferences are shaped by observation (Panizza, et al., 2021; Suzuki et al. 2016; Yu et al, 2021), although little-to-no work has examined whether contagion differs for relative vs. absolute preferences. Associative models predict that social contagion may be exaggerated in BPD (Ereira et al., 2018).… As a whole, humans are more susceptible to changing relative preferences more than selfish, absolute reward preferences, and this is disrupted in BPD.’

Psychometric and Intentional Attribution analysis

‘Childhood trauma, persecution, and poor mentalising in BPD are all predicted to disrupt one’s ability to change (Fonagy & Luyten, 2009).’

‘Prior work has also predicted that partner-participant preference disparity influences mental state attributions (Barnby et al., 2022; Panizza et al., 2021).’

I'm not sure I understand why the authors, after adding multiple comparison correction, now list two kinds of p-values. To me, this is misleading and precludes the point of multiple comparison corrections, I therefore recommend they report the FDR-adjusted p-values only. Likewise, if a corrected p-value is greater than 0.05 this should not be interpreted as a result.

We have now adjusted the exploratory results to include only the FDR corrected values in the text.

‘We assessed conditional psychometric associations with social contagion under the assumption of M3 for all participants. We conducted partial correlation analyses to estimate relationships conditional on all other associations and retained all that survived bootstrapping (5000 reps), permutation testing (5000 reps), and subsequent FDR correction. When not controlled for group status, RGPTSB and CTQ scores were both moderately associated with MZQ scores (RGPTSB r = 0.41, 95%CI: 0.23, 0.60, p[fdr]=0.043; CTQ r = 0.354 95%CI: 0.13, 0.56, p[fdr]=0.02). This was not affected by group correction. CTQ scores were moderately and negatively associated with shifts in individualistic reward preferences (; r = -0.25, 95%CI: -0.46, -0.04, p[fdr]=0.03). This was not affected by group correction. MZQ scores were in turn moderately and negatively associated with shifts in prosocial-competitive preferences () between phase 1 and 3 (r = -0.26, 95%CI: -0.46, -0.06, p[fdr]=0.03). This was diminished when controlled for group status (r = 0.13, 95%CI: -0.34, 0.08, p[fdr]=0.20). Together this provides some evidence that self-reported trauma and self-reported mentalising influence social contagion (Fig S11). Social contagion under M3 was highly correlated with contagion under M1 demonstrating parsimony of outcomes across models (Fig S12).

Prior work has predicted that partner-participant preference disparity influences mental state attributions (Barnby et al., 2022; Panizza et al., 2021). We tested parameter influences on explicit intentional attributions in Phase 2 while controlling for group status. Attributions included the degree to which they believed their partner was motived by harmful intent (HI) and self-interest (SI). According with prior work (Barnby et al., 2022), greater disparity of absolute preferences before learning was associated on a trend level with reduced attributions of SI (<= -0.23, p[fdr]=0.08), and greater disparity of relative preferences before learning exaggerated attributions of HI = 0.21, p[fdr]=0.08), but did not survive correction (Figure S4B). This is likely due to partners being significantly less individualistic and prosocial on average compared to participants (= -5.50, 95%HDI: -7.60, -3.60; = 12, 95%HDI: 9.70, 14.00); partners are recognised as less selfish and more competitive.’

Can the authors please elaborate why the algorithm proposed to be employed by BPD is more 'entropic', especially given both their self-priors and posteriors about partners' preferences tended to be more precise than the ones used by CON? As far as I understand, there's nothing in the data to suggest BPD predictions should be more uncertain. In fact, this leads me to wonder, similarly to what another reviewer has already suggested, whether BPD participants generate self-referential priors over others in the same way CON participants do, they are just less favourable (i.e., in relation to oneself, but always less prosocial) - I think there is currently no model that would incorporate this possibility? It should at least be possible to explore this by checking if there is any statistical relationship between the estimated θ_ppt^m and 〖p(θ〗_par |D^0).

Thank you for this opportunity to be clearer in our wording. We belief the reviewer is referring to this line in the discussion: ‘In either case, the algorithm underlying the computational goal for BPD participants is far higher in entropy and emphasises a less stable or reliable process of inference.’

We note in the revised Figure 2 panel E and in the results that those with BPD under M4 show insertion along absolute reward (they still expect diminished selfishness in others), but neutral priors over relative reward (around 0, suggesting expectations of neither prosocial or competitive tendencies of others). Thus, θ_ppt^m (self preference) and θ_par^m (other preference) are tightly associated for absolute, but not relative reward.

In our wording, we meant that whether under model M4 or M1, those with BPD either show a neutral prior over relative reward (M4) or a prior with large variance over relative reward (M1), showing expectations of difference between themselves and their partner. In both cases, expectation about a partner’s absolute reward preferences is diminished vs. CON participants. We have strengthened our language in the discussion to clarify this:

‘In either case, the algorithm underlying the computational goal for BPD participants is far higher in uncertainty, whether through a neutral central tendency (M4) or large variance (M1) prior over relative reward in phase 2, and emphasises a less certain and reliable expectation about others.’

To note, social contagion under M3 was highly correlated with contagion under M1 (see Fig S11). This provides some preliminary evidence that trauma impacts beliefs about individualism directly, whereas trauma and persecutory beliefs impact beliefs about prosociality through impaired trait mentalising" - I don't understand what the authors mean by this, can they please elaborate and add some explanation to the main text?

We have now clarified this in the text:

‘Together this provides some evidence that self-reported trauma and self-reported mentalising influence social contagion (Fig S11). Social contagion under M3 was highly correlated with contagion under M1 demonstrating parsimony of outcomes across models (Fig S12).’

I noted that at least some of the newly added references have not been added to the bibliography (e.g., Hitchcock et al. 2022).

Thankyou for noticing this omission. We have now ensured all cited works are in the reference list.

Reviewer 2:

The paper is not based on specific empirical hypotheses formulated at the outset, but, rather, it uses an exploratory approach. Indeed, the task is not chosen in order to tackle specific empirical hypotheses. This, in my view, is a limitation since the introduction reads a bit vague and it is not always clear which gaps in the literature the paper aims to fill. As a further consequence, it is not always clear how the findings speak to previous theories on the topic.’

As I wrote in the public review, however, I believe that an important limitation of this work is that it was not based on testing specific empirical hypotheses formulated at the outset, and on selecting the experimental paradigm accordingly. This is a limitation because it is not always clear which gaps in the literature the paper aims to fill. As a consequence, although it has improved substantially compared to the previous version, the introduction remains a bit vague. As a further consequence, it is not always clear how the findings speak to previous theories on the topic. Still, despite this limitation, the paper has many strengths, and I believe it is now ready for publication

Thank you for this further critique. We appreciate your appraisal that the work has improved substantially and is ready for publication. We nevertheless have opted to clarify our introduction and aprior predictions throughout the manuscript (please see response to Reviewer 1).

Reviewer 3:

Although the authors note that their approach makes "clear and transparent a priori predictions," the paper could be improved by providing a clear and consolidated statement of these predictions so that the results could be interpreted vis-a-vis any a priori hypotheses.

In line with comments from both Reviewer 1 and 2, we have clarified our introduction to make it clear what our aprior predictions and hypotheses are about our core aims and exploratory analyses (see response to Reviewer 1).

The approach of using a partial correlation network with bootstrapping (and permutation) was interesting, but the logic of the analysis was not clearly stated. In particular, there are large group (Table 1: CON vs. BPD) differences in the measures introduced into this network. As a result, it is hard to understand whether any partial correlations are driven primarily by mean differences in severity (correlations tend to be inflated in extreme groups designs due to the absence of observation in middle of scales forming each bivariate distribution). I would have found these exploratory analyses more revealing if group membership was controlled for.

Thank you for this chance to be clearer in our methods. We have now written a more direct exposition of this exploratory method:

‘Exploratory Network Analysis

To understand the individual differences of trait attributes (MZQ, RGPTSB, CTQ) with other-to-self information transfer () across the entire sample we performed a network analysis (Borsboom, 2021). Network analysis allows for conditional associations between variables to be estimated; each association is controlled for by all other associations in the network. It also allows for visual inspection of the conditional relationships to get an intuition for how variables are interrelated as a whole (see Fig S11). We implemented network analysis with the bootNet package in r using the ‘estimateNetwork’ function with partial correlations (Epskamp, Borsboom & Fried, 2018). To assess the stability of the partial correlations we further implemented bootstrap resampling with 5000 repetitions using the ‘bootnet’ function. We then additionally shuffled the data and refitted the network 5000 times to determine a ppermuted value; this indicates the probability that a conditional relationship in the original network was within the null distribution of each conditional relationship. We then performed False Discovery Rate correction on the resulting p-values. We additionally controlled for group status for all variables in a supplementary analysis (Table S4).’

We have also further corrected for group status and reported these results as a supplementary table, and also within the main text alongside the main results. We have opted to relegate Figure 4 into a supplementary figure to make the text clearer.

‘We explored conditional psychometric associations with social contagion under the assumption of M3 for all participants (where everyone is able to be influenced by their partner). We conducted partial correlation analyses to estimate relationships conditional on all other associations and retained all that survived bootstrapping (5000 reps), permutation testing (5000 reps), and subsequent FDR correction. When not controlled for group status, RGPTSB and CTQ scores were both moderately associated with MZQ scores (RGPTSB r = 0.41, 95%CI: 0.23, 0.60, p[fdr]=0.043; CTQ r = 0.354 95%CI: 0.13, 0.56, p[fdr]=0.02). This was not affected by group correction. CTQ scores were moderately and negatively associated with shifts in individualistic reward preferences (; r = -0.25, 95%CI: -0.46, -0.04, p[fdr]=0.03). This was not affected by group correction. MZQ scores were in turn moderately and negatively associated with shifts in prosocial-competitive preferences () between phase 1 and 3 (r = -0.26, 95%CI: -0.46, -0.06, p[fdr]=0.03). This was diminished when controlled for group status (r = 0.13, 95%CI: -0.34, 0.08, p[fdr]=0.20). Together this provides some evidence that self-reported trauma and self-reported mentalising influence social contagion (Fig S11). Social contagion under M3 was highly correlated with contagion under M1 demonstrating parsimony of outcomes across models (Fig S12).’

Discussion first para: "effected -> affected"

Thanks for spotting this. We have now changed it.

Add "s" to "participant: "Notably, despite differing strategies, those with BPD achieved similar accuracy to CON participant."

We have now changed this.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Generalising information from ourselves to others, and others to ourselves allows for both a dependable source of navigation and adaptability in interpersonal exchange. Disturbances to social development in sensitive periods can cause enduring and distressing damage to lasting healthy relationships. However, identifying the mechanisms of healthy exchange has been difficult. We introduce a theory of self-other generalisation tested with data from a three-phase social value orientation task – the Intentions Game. We involved individuals with (n=50) and without (n=53) a diagnosis of borderline personality disorder and assessed whether infractions to self-other generalisation may explain interpersonal (in)stability. Healthy controls initially used their preferences to predict others and were influenced by their partners, leading to self-other convergence. In contrast, individuals with borderline personality disorder maintained distinct self-other representations, generating a new neutral prior when learning about others. This allowed for equal predictive performance compared to controls despite reduced updating sensitivity. Furthermore, we explored theory-driven individual differences underpinning contagion. Overall, the findings provide a clear explanation of how self-other generalisation constrains and assists learning and how childhood adversity is associated with separation of internalised beliefs. Our model makes clear causal predictions about the mechanisms of social information generalisation concerning both joint and individual reward.

Revised version: Generalising information from ourselves to others, and others to ourselves allows for both a dependable source of navigation and adaptability in interpersonal exchange. Disturbances to social development in sensitive periods can cause enduring and distressing damage to lasting healthy relationships. However, identifying the mechanisms of healthy exchange has been difficult. We introduce a theory of self-other generalisation tested with data from a three-phase social value orientation task – the Intentions Game. We involved individuals with (n=50) and without (n=53) a diagnosis of borderline personality disorder and assessed whether infractions to self-other generalisation may explain prior findings of disrupted social learning and instability. Healthy controls initially used their preferences to predict others and were influenced by their partners, leading to self-other convergence. In contrast, individuals with borderline personality disorder maintained distinct self-other representations when learning about others. This allowed for equal predictive performance compared to controls despite reduced updating sensitivity. Furthermore, we explored theory-driven individual differences underpinning contagion. Overall, the findings provide a clear explanation of how self-other generalisation constrains and assists learning and how childhood adversity is associated with separation of internalised beliefs. Our model makes clear predictions about the mechanisms of social information generalisation concerning both joint and individual reward.

Introduction

Added in the revised version: Social animals have evolved sophisticated mechanisms for cooperation, exchanging information to enable both individual and group regulation (Emerson, 1956; Wheeler, 1911). In humans, such exchanges provide crucial insights about others as well as oneself, fostering the development of representations encoding self and others that is fundamental for adaptive social orientation and interaction. Disruptions to this process can impair the formation of stable social bonds and result in rigid interpersonal beliefs (Fairbairn, 1952; Young et al., 2006).

Added in the revised version: To effectively navigate what is sometimes called relational uncertainty, individuals generalize information from themselves to others (self-to-other transfer) and from others to themselves (other-to-self transfer). The relational self theory (Anderson & Chen, 2002) posits that when uncertain about others’ behaviours, people often rely on their own preferences as initial hypotheses, a process termed self-insertion (Allport, 1924; Kreuger & Clement, 1994). Conversely, when uncertain about their own states, individuals use external social cues to adjust their self-representations, a process known as social contagion (Deutsch & Gerard, 1955; Toelch & Dolan, 2015; Moutoussis et al, 2016). This bidirectional generalization has been widely observed across various domains including economic decision-making, morality, and social group adaptation (Devaine & Daunizeu, 2017; Garvert et al., 2015; Panizza et al., 2021; Suzuki et al., 2016; Yu et al., 2021).

Added in the revised version: Understanding healthy interpersonal dynamics can be clarified by examining their disruptions. Individuals with Borderline Personality Disorder (BPD) provide a compelling case study due to their profound interpersonal instability, emotional dysregulation, and heightened sensitivity to social contexts (Gunderson et al., 2018). BPD is strongly associated with adverse childhood experiences (Afifi et al., 2011), impaired mentalizing abilities (Fonagy & Luyten, 2009), and maladaptive representations of self and other (Hanegraaf et al., 2021). This phenotype has been explained through disrupted and unstable social inference during observation (Henco et al., 2020; Hula et al., 2018, Story et al., 2024a; Siegel et al., 2020). However, the precise mechanisms linking disrupted social cognition in BPD remain elusive, particularly regarding whether individuals with BPD differ specifically in their use of self-to-other and other-to-self information transfer.

Added in the revised version: This paper seeks to address this gap by testing explicitly how disruptions in self-other generalization processes may underpin interpersonal disruptions observed in BPD. Specifically, our hypotheses were: (i) healthy controls will demonstrate evidence for both self-insertion and social contagion, integrating self and other information during interpersonal learning; and (ii) individuals with BPD will exhibit diminished self-other integration, reflected in stronger evidence for observations that assume distinct selfother representations.

Added in the revised version: We tested these hypotheses by designing a dynamic, sequential, three-phase Social Value Orientation (Murphy & Ackerman, 2014) paradigm—the Intentions Game—that would provide behavioural signatures assessing whether BPD differed from healthy controls in these generalization processes (Figure 1A). We coupled this paradigm with a lattice of models (M1-M4) that distinguish between self-insertion and social contagion (Figure 1B), and performed model comparison:

Added in the revised version: M1. Both self-to-other (self-insertion) and other-to-self (social contagion) occur before and after learning

Added in the revised version: M2. Self-to-other transfer only occurs

Added in the revised version: M3. Other-to-self transfer only occurs

Added in the revised version: M4. Neither transfer process, suggesting distinct self-other representations

Added in the revised version: We additionally ran exploratory analysis of parameter differences and model predictions between groups following from prior work demonstrating changes in prosociality (Hula et al., 2018), social concern (Henco et al., 2020), belief stability (Story et al., 2024a), and belief updating (Story, 2024b) in BPD to understand whether discrepancies in self-other generalisation influences observational learning. By clearly articulating our hypotheses, we aim to clarify the theoretical contribution of our findings to existing literature on social learning, BPD, and computational psychiatry.

Removed in the revised version: Social animals have evolved sophisticated mechanisms for cooperation. To coordinate, individuals exchange information, enabling independent and group regulation (Emerson, 1956; Wheeler, 1911). Interaction provides the self with insights about others and, simultaneously, about its own state within the environment. In humans, this psychological differentiation begins in utero (Ciaunica et al., 2021) and continues into adulthood. This prolonged differentiation facilitates the distinction between self and others, allowing grounded social orientation and flexible adaptation. Early psychological theorists highlighted the importance of healthy social exchange in humans for integrating constructive relational blueprints, noting that disruptions during sensitive periods can impair the formation of safe social bonds (e.g., Fairbairn, 1952; 1994) and foster less adaptable interpersonal beliefs (Young et al., 2006).

Removed in the revised version: At the individual level, humans must manage information exchange to navigate and reduce relational uncertainty. When faced with external uncertainty about others’ characteristics, prior knowledge can swiftly and effectively guide predictions. Indeed, to reduce uncertainty about others, theories of the relational self (Anderson & Chen, 2002) suggest that the self is the most extensive and well-grounded representation available, leading to a readily accessible initial belief that can be projected or integrated when learning about others (self-insertion; Allport, 1924; Kreuger & Clement, 1994). When uncertainty arises regarding one’s own internal state, external cues can provide anchoring information to support self-through-other calibration. Therefore, individuals can generalize information from others to themselves. contagion has been observed across economic (Suzuki et al., 2016), intertemporal (Garvert et al., 2015), effort discounting (Devaine & Daunzieu, 2017), and moral choice behaviour (Yu et al., 2021). Social contagion is hypothesised to facilitate adaptation to social groups (Deutsch & Gerard, 1955; Toelch & Dolan, 2015), and as a crucial component of interpersonal cohesion that relies on trust (Frith & Frith, 2012).

Removed in the revised version: Insights into healthy self-other relations can be gained from their frequent absence, a common and distressing reality for those with a diagnosis of Borderline Personality Disorder (BPD). BPD is characterized by interpersonal sensitivity and instability, emotional dysregulation, impulsivity, paranoia, and severe fear of abandonment (e.g., Gunderson et al., 2018; Euler et al., 2021). It is strongly associated with early childhood adversity, such as psychological, physical, and sexual abuse or neglect (Afifi et al., 2011; Bateman et al., 2023), and inconsistent parenting experiences (Crawford et al., 2009), particularly during sensitive developmental periods (Fonagy & Bateman, 2008). Early adversity interacts with pre-existing variations in stress-homeostasis mechanisms, exacerbating interpersonal disruptions (Pratt et al., 2017), underpinned by difficulties in mentalizing, leading to disruptions in social learning (e.g., Fonagy & Luyten, 2009; Nolte et al., 2023) and representations of self and other (Hanegraaf et al., 2021). The mechanisms connecting early life adversity and relational instability are elusive, making it harder to find a path to rehabilitation.

Removed in the revised version: Critical questions persist concerning both typical and disrupted interaction: What computational mechanisms support efficient generalisation from self-to-other and other-to-self? Do processes operate globally or locally based on outcome (e.g. individual vs. joint reward)? Can such formal models explain observations of social learning changes in diagnoses like BPD? Understanding the processes and representations humans use in social exchange has broad implications for theories of development, group cohesion, reputation management, interpersonal synchrony, and the breakdown of social bonds.

Removed in the revised version: Formal models of healthy (non-clinical) exchange posit forms of interaction between self and other representations which can reduce uncertainty during learning (whilst potentially introducing biases). As such, self-insertion accounts demonstrate that integrating self-preferences into initial predictions about others causally increases reaction times and can increase predictive accuracy when predictions are correct (Tarantola et al., 2017; Barnby et al., 2022). Social contagion has been framed as convergence to a common latent state shared by oneself and another, calibrated by self and other uncertainty (Moutoussis et al., 2016; Thomas et al., 2022). By comparison, in mental ill-health, computational models have almost exclusively probed social processes during observational learning without fully integrating self and other into future beliefs. The BPD phenotype has been associated with a potential over-reliance on social versus internal cues (Henco et al., 2020), ‘splitting’ of social latent states that encode beliefs about others (Story et al., 2023), negative appraisal of interpersonal experiences with heightened self-blame (Mancinelli et al., 2024), inaccurate inferences about others’ irritability (Hula et al., 2018), and reduced belief adaptation in social learning contexts (Siegel et al., 2020). Associative models have also been adapted to characterize ‘leaky’ self-other reinforcement learning (Ereira et al., 2018), finding that those with BPD overgeneralize (leak updates) about themselves to others (Story et al., 2024). Altogether, there is currently a gap in the direct causal link between insertion, contagion, and learning (in)stability.

Removed in the revised version: Our present work sought to achieve two primary goals. 1. Extend prior computational theories to formalise and test the interrelation between self-insertion and social contagion in learning and choice to probe self-other generalisation, and 2. Test whether previous computational findings of differences in social learning in people with BPD can be explained by disruptions to self-other generalisation. Here, our formal hypotheses concerning self-other generalisation (Barnby et al., 2023) make clear a priori predictions about its relationship with observational learning that can be empirically interrogated. Exploring these factors within models can uncover algorithmic biases that healthy volunteers and people with BPD either do or do not share. We accomplish these goals using a dynamic, sequential economic paradigm called the Intentions Game. The Intentions Game builds upon the well-known Social Value Orientation framework (Murphy & Ackerman, 2011), which captures innate motivational variation in reward allocation to self and other.

Removed in the revised version: We discover that healthy participants employ a mixed process of self-insertion and contagion to predict and align with the beliefs of their partners. In contrast, age, sex, and education, and social deprivation matched individuals with BPD exhibit distinct, disintegrated representations of self and other, despite showing similar average accuracy in their predictions about partners. Our model and data suggest that the previously observed computational characteristics in BPD arise from a failure of information transfer, first from self to other, and then from other to self, an under rather than over generalisation. By integrating separate computational findings, we provide a foundational model and a concise, dynamic paradigm to investigate uncertainty, generalisation, and learning in social interactions.

Results

Healthy participants (CON; n=53) and participants diagnosed with BPD (n=50), matched on age, gender, education, and social deprivation indices (Table 1), were invited to participate in a three-phase social value orientation paradigm—the Intentions Game (Figure 1A)—with virtual partners. In phase 1, participants made forced choices between two options for splitting points with an anonymous partner. In phase 2 participants learned to predict the decisions of a new anonymous partner using the same forced-choice setup, receiving feedback on the accuracy of their successive predictions. Notably, using a novel server architecture (Burgess et al., 2023), partners in phase 2 were configured to be approximately 50% different from the participants in terms of their choices, ensuring that all participants had to learn about their partners. Phase 3 mirrored the first, with participants informed that they were matched with a third anonymous partner, unrelated to those in phases 1 and 2. Detailed descriptions of the task can be found in the methods section and Figure 1. All participants also self-reported their trait paranoia, childhood trauma, trust beliefs, and trait mentalizing (see Methods).

Psychometric and Behavioural Results

Participants with BPD, compared to CON, retrospectively reported significant childhood trauma, epistemic disruptions (including mistrust and credulity), elevated referential and persecutory beliefs, and demonstrated ineffective trait mentalizing (Table 1). The groups did not differ in trait measures of certainty regarding self and others’ mental states, nor in epistemic trust.

We analysed the ‘types’ of choices participants made in each phase (Supplementary Table 1). The interpretation of a participant’s choice depends on both values in a choice. For example, a participant could make prosocial (self=5; other=5) versus individualistic (self=10; other=5) choices, or prosocial (self=10; other=10) versus competitive (self=10; other=5) choices. There were 12 of each pair in phases 1 and 3 (individualistic vs. prosocial; prosocial vs. competitive; individualistic vs. competitive).

In phase 1, both CON and BPD participants made prosocial over competitive choices with similar frequency (CON=9.67[3.62]; BPD=9.60[3.57]; t=-0.11, p=0.91). However, CON participants made significantly fewer prosocial choices when individualistic choices were available (CON=2.87[4.01]; BPD=5.22[4.54]; t=2.75, p=0.007). Both groups favoured individualistic over competitive choices with similar frequency (CON=11.03[1.95]; BPD=10.34[2.63]; t=-1.52, p=0.13). For a reaction time assessment see Supplementary Text 1).

Each group showed good predictive accuracy (CON=77.2%[13.9%]; BPD=72.7%[15.6%]). There was no difference in overall predictive accuracy between BPD and CON (linear estimate=2.44, 95%CI: −0.67, 5.54; t=1.56; p=0.12), nor on a trial-by-trial basis (linear estimate=0.26, 95%CI: −0.06, 0.59; z=1.61, p=0.11). All participants showed an effect of time on accuracy, such that participants became more accurate in predicting their partner over the course of phase 2 (linear estimate=0.013, 95%CI: 0.008, 0.017; z=6.01; p<0.001). Server matching between participant and partner in phase 2 was successful, with participants being approximately 50% different to their partners with respect to the choices each would have made on each trial in phase 2 (mean similarity=0.49, SD=0.12).

In phase 3, both CON and BPD participants continued to make equally frequent prosocial versus competitive choices (CON=9.15[3.91]; BPD=9.38[3.31]; t=-0.54, p=0.59); CON participants continued to make significantly less prosocial versus individualistic choices (CON=2.03[3.45]; BPD=3.78 [4.16]; t=2.31, p=0.02). Both groups chose equally frequent individualistic versus competitive choices (CON=10.91[2.40]; BPD=10.18[2.72]; t=-0.49, p=0.62).

Task and Model Space.

(A) Participants were invited to play a three-phase, repeated social value orientation paradigm—the Intentions Game—with virtual partners. Phase 1 of the Intentions Game lasted 36 trials and asks participants to make a forced choice between two options as to how to split points with an anonymous virtual partner. An example of a prosocial-individualistic pair of options could be (self=5, other=5) or (self=10, other=5) – if the participant chooses option 1 they could be viewed as less individualistic and more prosocial as the outcomes to the other do not change, but the self would earn less. In phase 2, lasting 54 trials, participants were asked to predict the decisions of a new anonymous partner using the same two-forced choice set-up and the same option pairs; participants were given feedback on whether they were correct or incorrect in their prediction. We used Amazon Web Services to create a novel server architecture to match participants and (virtual) partners (Burgess et al., 2023). Partners in phase 2 were matched to be approximately 50% different from the participant with respect to their choices in phase 1 to ensure all participants needed to learn about their phase 2 partner, and to provide a mechanism to examine whether beliefs about partners had an effect on the self. Phase 3 was identical to phase 1, although participants were informed that they were matched with a third anonymous partner, unconnected to the partners in phase 1 and 2. At the end of the game, if participants collected over 1000 points overall, they were entered into a lottery to win a bonus. (B) We created four models that may explain the data and to test theories of social generalization. Model M1 assumes participants are subject to both self-insertion and social-contagion, that is, participants used their own preferences as a prior about their partner in phase 2, and partner behaviour subsequently influenced participant’s preferences in phase 3. Model M4 assumes participants are subject to neither self-insertion nor social contagion, instead forming a novel prior around the phase 2 partner rather than using their own preferences and failing to be influenced by their partner after observation. Models M2 and M3 suggest participants are only explained by either self-insertion or social-contagion, not both. (C) We assume that participants choices in phase 1 are governed by both a median ( θpptm ) and standard deviation ( θpptσ ). Participants insert their median preferences ( θpptm ) into their prior beliefs over their partner in phase 2, but with a different standard deviation to allow for flexibility and learning ( θrefσ ). The combination of the prior and posterior belief uncertainty about the partner ( θrefσ ; θparσ ), the precision participants have over their own preferences ( θpptσ ), and the median posterior of the participant and partner ( θpptm ; θparm ) form the new median and standard deviation over participant preferences in phase 3 ( θ′ppt ). (D) In contrast to M1, M4 generates a new central tendency over the partner in phase 2 ( θ¯parm ) which disconnects participant preferences and prior beliefs. M4 also assumes that the same parameters that generated participant choices in phase 1 also generate choices in phase 3. (E) Simulating our model demonstrates how different combinations of α (preferences for absolute self-reward) and β (preferences for relative reward; prosocial-competitiveness) lead to changes in the discrepancy of value between participants and partners (left panel). We also show how increasing uncertainty over self-beliefs, and higher precision over partners, causally draws participants more toward the beliefs of their partner in phase 3 and increases their precision over their phase 3 beliefs (Moutoussis et al., 2016).

Computational Analysis

Over all three phases, we assumed participants and their partners used a Fehr-Schmidt utility function (Fehr & Schmidt, 1999) to calculate the utility of two options (Uα,β={ Uα,β1,Uα,β2 }) , based on the joint rewards available for both the participant, Rppt = { rppt1,rppt2 } and their partner, Rpar = { rppt1,rppt2 } . The utility of each option was weighted based on absolute-reward gain α (how much participants care about self-earnings, irrespective of the other) and relative-reward β along a prosocial-competitive axis (how much participants care about equality of outcomes); β < 0 is prosocial, whereas β > 0 is competitive.

We then constructed four models to explain how participants used their own preferences ( θpptm = { αpptm , βpptm }) and uncertainty over these preferences ( θpptσ = { αpptσ , βpptσ }) to predict and learn about the preferences of their partner (θpar; Figure 1B; see methods). Model M1 (Figure 1C) suggests that participants initially use their own preferences as a prior belief about their partner (self-insertion), which is gradually diminished during the learning process in phase 2. M1 also posits that, following learning, the inferred beliefs about a partner will influence participants’ own preferences, making them more similar to their partner’s preferences following observation (social contagion). According to this model, participants shift towards their partner based on their uncertainty about self and others (Figure 1E): greater uncertainty over self-preferences and increased precision in representing the other cause stronger social contagion effects.

Earlier version: Model M4 (Figure 1D), on the other hand, suggests that participants do not engage in these generalization processes: predictions about others are not grounded in the self, and observing others does not alter self-preferences. Models M2 and M3 allow for either self-insertion or social contagion to occur independently. Consistent with prior research, we also constructed a model that assumes the same insertion and contagion processes as M1, but along a single prosocial-competitive axis (‘Beta model’; Barnby et al., 2022). The ‘Beta model’ is equivalent to M1 in its causal architecture (both self-insertion and social contagion are hypothesized to occur) but differs in its utility function: participants might only consider a single dimension of relative reward allocation, which is typically emphasized in previous studies (e.g., Hula et al., 2018).

Revised version: Model M4 (Figure 1D), on the other hand, suggests that participants do not engage in these generalization processes: predictions about others are not grounded in the self, and observing others does not alter self-preferences. Models M2 and M3 allow for either self-insertion or social contagion to occur independently. Consistent with prior research, we also constructed a model that assumes the same insertion and contagion processes as M1, but along a single prosocial-competitive axis (‘Beta model’; Barnby et al., 2022). The ‘Beta model’ is equivalent to M1 in its architecture (both self-insertion and social contagion are hypothesized to occur) but differs in its utility function: participants might only consider a single dimension of relative reward allocation, which is typically emphasized in previous studies (e.g., Hula et al., 2018).

All computational models were fitted and compared using a Hierarchical Bayesian Inference (HBI) algorithm which allows hierarchical parameter estimation while assuming random effects for group and individual model responsibility (Piray et al., 2019; see Methods for more information). We report individual and group-level model responsibility, in addition to exceedance probabilities between-groups to assess model dominance.

Model Comparison – BPD Participants Hold Disintegrated Self-Other Beliefs

Added in the revised version: We found that CON participants were best fit at the group level by M1 (Frequency = 0.59, Exceedance Probability = 0.98), whereas BPD participants were best fit by M4 (Frequency = 0.54, Exceedance Probability = 0.86; Figure 2A). This suggests CON participants are best fit by a model that fully integrates self and other when learning, whereas those with BPD are best explained as holding disintegrated and separate representations of self and other that do not transfer information back and forth.

Earlier version: We found that CON participants were best fit at the group level by M1 (Frequency = 0.59, Exceedance Probability = 0.98), whereas BPD participants were best fit by M4 (Frequency = 0.54, Exceedance Probability = 0.86; Figure 2A). We first analyse the results of these separate fits. Later, in order to assuage concerns about drawing inferences from different models, we examined the relationships between the relevant parameters when we forced all participants to be fit to each of the models (in a hierarchical manner, separated by group). In sum, our model comparison is supported by convergence in parameter values when comparisons are meaningful (see Supplementary Materials). We refer to both types of analysis below.

Revised version: We first explore parameters between separate fits (see Methods). Later, in order to assuage concerns about drawing inferences from different models, we examined the relationships between the relevant parameters when we forced all participants to be fit to each of the models (in a hierarchical manner, separated by group). In sum, our model comparison is supported by convergence in parameter values when comparisons are meaningful (see Supplementary Materials). We refer to both types of analysis below.

Generative Accuracy and Recovery

Earlier version: We simulated data for each participant using their individual parameters from the winning model within each group and refitted our models using this simulated data. Model comparison yielded very similar results (Figure 3A): CON synthetic participants best fit at the group level by M1 (Frequency = 0.58, Exceedance Probability = 0.98) and BPD synthetic participants best fit by M4 (Frequency = 0.57, Exceedance Probability = 0.85). The simulated data closely matched the actions of participants across all three phases (median accuracy = 0.8, SD = 0.12). In phase 2, the model-predicted total correct scores were not significantly different from observed scores (Figure 3E). Both model responsibility and common parameters within each dominant model were significantly associated (model confusion p = 0.46–0.97, p < 0.001; parameter recovery p = 0.70–0.94, p < 0.001; Figure 3C). βpptm , preferences in phase 1 were negatively correlated with prosocial versus competitive choices (r=-0.77, p<0.001) and individualistic versus competitive choices (r=-0.59, p<0.001); αpptm was positively correlated with individualistic versus competitive choices (r=0.53, p<0.001) and negatively correlated with prosocial versus individualistic choices (r=-0.69, p<0.001). Given the very good to excellent performance of the models, we analysed individual parameters and simulations between groups.

Revised version: We simulated data for each participant using their individual parameters from the winning model within each group and refitted our models using this simulated data. Model comparison yielded very similar results (Figure 3A): CON synthetic participants best fit at the group level by M1 (Frequency = 0.58, Exceedance Probability = 0.98) and BPD synthetic participants best fit by M4 (Frequency = 0.57, Exceedance Probability = 0.85). The simulated data closely matched the actions of participants across all three phases (median accuracy = 0.8, SD = 0.12). In phase 2, the model-predicted total correct scores were not significantly different from observed scores (Figure 3E). Both model responsibility and common parameters within each dominant model were significantly associated (model confusion ρ = 0.46–0.97, p < 0.001; parameter recovery ρ = 0.70–0.94, p < 0.001; Figure 3C). βpptm , preferences in phase 1 were negatively correlated with prosocial versus competitive choices (r=-0.77, p<0.001) and individualistic versus competitive choices (r=-0.59, p<0.001); αpptm was positively correlated with individualistic versus competitive choices (r=0.53, p<0.001) and negatively correlated with prosocial versus individualistic choices (r=-0.69, p<0.001). Given the very good to excellent performance of the models, we analysed individual parameters and simulations between groups.

Beliefs between groups and within phases.

(A) We used randomeffects hierarchical model fitting and comparison to jointly estimate group level and individual level parameters based on real data from participants (Piray et al., 2021). CON participants were best fit by M1, whereas BPD participants were best fit by M4 on a group level. Looking within each model by simulating the beliefs of each participant reveals that – as expected – CON participants use the median of their self-preferences (black distribution) as a basis for their prior beliefs about partners (light orange distribution), and that the precision of their posterior beliefs about partners (dark orange distribution) and the precision of their own self preferences leads to a shifted model of the self (grey distribution). BPD participants on the other hand have a disintegrated prior over their partner which is not subject to their own self representation. Likewise, there is no change in self-preferences following learning, and thus an absence of the light grey distribution. For illustration, we focus on beliefs over relative preferences (β) and use real individual participants as exemplars for illustration. (B) Across models we extracted the common parameters that generate the behaviour of both CON and BPD participants – that is, their median and standard deviation over both α (absolute reward preferences) and β (relative value preferences), the flexibility over participants’ prior beliefs about their partners over each dimension, and the absolute change in posterior beliefs in phase 2 over each dimension ( Δαparm ; Δβparm ). Using hierarchical Bayesian t-tests we demonstrated the mean difference in parameter values between groups. Purple values lower than 0 indicate that the BPD participants had significantly smaller parameter values. Here we find that BPD participants were less individualistic, equally prosocial, and more certain about their self-preferences. BPD participants were also less flexible over their beliefs about a partner’s absolute reward preferences and updated their beliefs less across the board. (C) Examining participants under a blanket assumption that participants in both BPD and CON groups were influenced by their partner (Model M3) revealed that BPD participants were significantly less influenced by their partner across the board, both with respect to their phase 3 median and standard deviation of beliefs. Kruskal-Wallis tests were used between groups within the visualisation. *=p<0.05, **=p<0.01, ***=p<0.001, ****=p<0.0001. (D) We also calculated the Kullback-Leibler divergence (DKL) of beliefs between each trial (t-1 vs t) on each trial during phase 2. We observed three things: 1. All participants display larger belief updates initially, 2. all participants ‘cool off’ in their belief updating over the course of phase 2, and 3. BPD participants update beliefs significantly less throughout the course of phase 2 vs. CON participants. (E) Examining the shift in central tendencies across both groups demonstrates that between phase 1 and the start of phase 2, BPD participants shifted their central tendency only over beliefs about the relative reward preferences of partners’. They held lower expectations about a partner’s absolute reward preferences consistent with their own preferences.

Phase 1 – BPD Participants Are More Certain About Themselves

We first examined self-representations of participants in phase 1. CON participants (under model M1) and BPD participants (under M4) were equally prosocial (CON mean[ βpptm ] = −7.50; BPD mean[ βpptm ] = −6.59; Δμ[ βpptm ] = 0.92, 95%HDI: −1.24, 3.12) – both groups valued equal allocation of reward between themselves and their partners. BPD participants had lower preferences for earning higher absolute rewards (CON mean[ αpptm ]=18.41; BPD mean[ αpptm ] = 10.57; Δμ[ αpptm ] = −7.83, 95%HDI: −11.06, −4.75). BPD participants were also more certain about both self-preferences for absolute and relative reward (Δμ[ αpptσ ] = −0.89, 95%HDI: −1.01, −0.75; Δμ[ βpptσ ] = −0.32, 95%HDI: −0.60, −0.04) versus CON participants (Figure 2B).

These differences were replicated when considering parameters between groups when we fit all participants to the same models (M1-M4; see Table S2).

Phase 2 – BPD Participants Use Disintegrated and Neutral Priors

We next assessed how participants generated their prior beliefs about a partner in phase 2. CON participants were best fit by M1 which assumes the same median belief participants use in phase 1 is identical to their median prior belief about their partners. In contrast, BPD participants were best fit by M4 and so generated a new median prior belief about their partners. Assessing by individual models show this was driven by expectations about a partner’s prosocial-competitive preferences (relative reward; see Table S2).

Earlier version: In BPD participants, only new beliefs about the relative reward preferences of partners differed (see Fig 2E) – new median priors were larger than median preferences in phase 1 (mean[ βparm ] = −0.47; Δμ[ βpptm – βparm ] = −6.10, 95%HDI: −7.60, −4.60). BPD priors about their partner’s relative preferences were also centred closely around 0 (Δμ[0 − βparm ] = −0.39, 95%HDI: −0.77, −0.05), suggesting that BPD participants entered into the interaction with very neutral priors about their partner’s preferences for relative reward.

Revised version: Prior work predicts those with BPD should focus more intently on social information rather than private information that only concerns one party (Henco et al., 2020). In BPD participants, only new beliefs about the relative reward preferences – mutual outcomes for both player - of partners differed (see Fig 2E): new median priors were larger than median preferences in phase 1 (mean[ βparm ] = −0.47; Δμ[ βpptm – βparm ] = −6.10, 95%HDI: −7.60, −4.60). BPD priors about their partner’s relative preferences were also centred closely around 0 (Δμ[0 − βparm ] = −0.39, 95%HDI: −0.77, −0.05), suggesting that BPD participants entered into the interaction with very neutral priors about their partner’s preferences for relative reward.

Earlier version: BPD participants were equally flexible around their prior beliefs about a partner’s relative reward preferences (Δμ[ βparref ] = −1.60, 95%HDI: −3.42, 0.23), and were less flexible around their beliefs about a partner’s absolute reward preferences (Δμ[ βparref ]=- 4.09, 95%HDI: −5.37, −2.80), versus CON (Figure 2B). Median belief change (from priors to posteriors) in phase 2 was lower in BPD versus CON (Δμ[ αparm ] = −5.53, 95%HDI: −7.20, −4.00; Δμ[Δ βparm ] = −10.02, 95%HDI: −12.81, −7.30). Posterior beliefs about partner were more precise in BPD versus CON (Δμ[ αparσ ] = −0.94, 95%HDI: − 1.50, −0.45; Δμ[ βparσ ] = −0.70, 95%HDI: −1.20, −0.25). This is unsurprising given the disintegrated priors of the BPD group in M4, meaning they need to ‘travel less’ in state space. Even under assumptions of M1-M4 for both groups, BPD versus CON showed smaller posterior standard deviation in phase 2 over after learning (see Table T2). These results converge to suggest those with BPD form rigid posterior beliefs about prosocial-competitiveness.

Revised version: Models of moral preference learning (Story et al., 2024a) predict that BPD vs non-BPD participants have more rigid beliefs about their partners. We found that BPD participants were equally flexible around their prior beliefs about a partner’s relative reward preferences (Δμ[ βparref ] = −1.60, 95%HDI: −3.42, 0.23), and were less flexible around their beliefs about a partner’s absolute reward preferences (Δμ[ βparref ]=- 4.09, 95%HDI: −5.37, −2.80), versus CON (Figure 2B). Median belief change (from priors to posteriors) in phase 2 was lower in BPD versus CON (Δμ[ αparm ] = −5.53, 95%HDI: −7.20, −4.00; Δμ[Δ βparm ] = −10.02, 95%HDI: −12.81, −7.30). Posterior beliefs about partner were more precise in BPD versus CON (Δμ[ αparσ ] = −0.94, 95%HDI: − 1.50, −0.45; Δμ[ βparσ ] = −0.70, 95%HDI: −1.20, −0.25). This is unsurprising given the disintegrated priors of the BPD group in M4, meaning they need to ‘travel less’ in state space. Even under assumptions of M1-M4 for both groups, BPD versus CON showed smaller posterior standard deviation in phase 2 over after learning (see Table S2). These results converge to suggest those with BPD form rigid posterior beliefs.

Earlier version: We checked that these conclusions about self-insertion did not depend on the different models, we found that βparref under M1 and M2 were consistently larger in BPD versus CON. This supports the notion that new central tendencies for BPD participants in phase 2 were required for expectations about a partner’s relative reward (see Fig S10 & Table S2). αparref and βparref parameters under assumptions of M1 and M2 were strongly correlated with median change in belief between phase 1 and 2 under M3 and M4, suggesting convergence in model outcome (Fig S11).

Revised version: We checked that conclusions about self-insertion did not depend on the different models, we found that βparref under M1 and M2 were consistently larger in BPD versus CON. This supports the notion that new central tendencies for BPD participants in phase 2 were required for expectations about a partner’s relative reward (see Fig S10 & Table S2). αparref and βparref parameters under assumptions of M1 and M2 were strongly correlated with median change in belief between phase 1 and 2 under M3 and M4, suggesting convergence in model outcome (Fig S11).

Analysing belief updating on a more granular trial-by-trial basis using M1 for CON and M4 for BPD revealed preference type and between-group differences in belief refinement over the course of phase 2 (Figure 2D). We examined this by analysing the Kullback-Leibler divergence (DKL) – expected informational surprise - on each trial in Phase 2.

Earlier version: Across both groups and belief types informational surprise reduced over time (linear estimate[DKL] = −0.007, 95%CI: −0.008, −0.005; t = −7.60, p < 0.001). Beliefs about a partner’s relative reward preferences were updated more than absolute reward preferences (linear estimate= 0.54, 95%CI: 0.47, 0.62; t = 14.00, p < 0.001). These interacted, updating over relative vs. absolute beliefs reduced over the course of phase 2 (linear estimate = −0.013, 95%CI: −0.015, −0.011; t = −10.81, p < 0.001). These findings were supported under M1-M4 only assumptions (see Table T3).

Revised version: Across both groups and belief types informational surprise reduced over time (linear estimate[DKL] = −0.007, 95%CI: −0.008, −0.005; t = −7.60, p < 0.001). Beliefs about a partner’s relative reward preferences were updated more than absolute reward preferences (linear estimate= 0.54, 95%CI: 0.47, 0.62; t = 14.00, p < 0.001). These interacted, updating over relative vs. absolute beliefs reduced over the course of phase 2 (linear estimate = −0.013, 95%CI: −0.015, −0.011; t = −10.81, p < 0.001). These findings were supported under M1-M4 only assumptions (see Table S3).

Earlier version: BPD informational surprise is restricted over beliefs about absolute reward versus CON. CON participants remained more flexible than BPD participants along both types of preference (linear estimate [DKL( αparm )] = 0.40, 95%CI: 0.29, 0.51, t = 7.18, p < 0.001; linear estimate [DKL( βparm )] = 0.17, 95%CI: 0.29, 0.51, t=3.06, p=0.002). There was a group by time interaction, such that the difference between groups, and the magnitude of belief updates decreased over time (linear estimate [DKL( αparm )] = −0.009, 95%CI:-0.012,-0.006, t = −5.30, p < 0.001; linear estimate [DKL( βparm )] = −0.004, 95%CI:-0.008,-0.001, t = −2.78, p = 0.005) – CON participants and BPD participants eventually converged to an equivalent updating schedule (Figure 2D).

Revised version: BPD informational surprise is consistently restricted over beliefs about absolute reward versus CON. CON participants remained more flexible than BPD participants along both types of preference (linear estimate [DKL( αparm )] = 0.40, 95%CI: 0.29, 0.51, t = 7.18, p < 0.001; linear estimate [DKL( βparm )] = 0.17, 95%CI: 0.29, 0.51, t=3.06, p=0.002). There was a group by time interaction, such that the difference between groups, and the magnitude of belief updates decreased over time (linear estimate [DKL( αparm )] = −0.009, 95%CI:-0.012,-0.006, t = −5.30, p < 0.001; linear estimate [DKL( βparm )] = −0.004, 95%CI:-0.008,-0.001, t = −2.78, p = 0.005) - CON participants and BPD participants eventually converged to an equivalent updating schedule (Figure 2D).

Assessing this same relationship under M1- and M2-only assumptions reveals a replication of this group effect for absolute reward, but the effect is reversed for relative reward (see Table S3). This accords with the context of each model, where under M1 and M2, BPD participants had larger phase 2 prior flexibility over relative reward (leading to larger initial surprise), which was better accounted for by a new central tendency under M4 during model comparison. When comparing both groups under M1-M4 informational surprise over absolute reward was consistently restricted in BPD (Table S3), suggesting a diminished weight of this preference when forming beliefs about an other.

We explored how beliefs and choices were associated with reaction times, showing that belief updates and reaction times were coupled over the course of phase 2 and related to participant-partner similarity (Figure S9).

Phase 3 – BPD Participants Are Less Influenced by Partners

Earlier version: In the dominant model for the BPD group—M4—participants are not influenced in their phase 3 choices following exposure to their partner in phase 2. To further confirm this we also analysed absolute change in median participant beliefs between phase 1 and 3 under the assumption that M1 and M3 was the dominant model for both groups (that allow for contagion to occur). This analysis aligns with our primary model comparison using M1 for CON and M4 for BPD (Figure 2C). CON participants altered their median beliefs between phase 1 and 3 more than BPD participants (M1: linear estimate = 0.67, 95%CI: 0.16, 1.19; t = 2.57, p = 0.011; M3: linear estimate = 1.75, 95%CI: 0.73, 2.79; t = 3.36, p < 0.001). Relative reward was overall more susceptible to contagion versus absolute reward (M1: linear estimate = 1.40, 95%CI: 0.88, 1.92; t = 5.34, p<0.001; M3: linear estimate = 2.60, 95%CI: 1.57, 3.63; t = 4.98, p < 0.001). There was an interaction between group and belief type under M3 (M3: linear estimate = 2.13, 95%CI: 0.09, 4.18, t = 2.06, p=0.041) but not M1. There was a main effect of belief type on precision under M3 (linear estimate = 0.47, 95%CI: 0.07, 0.87, t = 2.34, p = 0.02) but not M1; relative reward preferences became more precise across the board. Derived model estimates of preference change between phase 1 and 3 strongly correlated between M1 and M3 along both belief types (see Table S2 and Fig S11).

Revised version: Prior work predicts that human economic preferences are shaped by observation (Panizza, et al., 2021; Suzuki et al. 2016; Yu et al, 2021). Associative models also predict that social contagion may be exaggerated in BPD (Story et al., 2024b). In the dominant model for the BPD group—M4—participants are not influenced in their phase 3 choices following exposure to their partner in phase 2. To further confirm this we also analysed absolute change in median participant beliefs between phase 1 and 3 under the assumption that M1 and M3 was the dominant model for both groups (that allow for contagion to occur). This analysis aligns with our primary model comparison using M1 for CON and M4 for BPD (Figure 2C). CON participants altered their median beliefs between phase 1 and 3 more than BPD participants (M1: linear estimate = 0.67, 95%CI: 0.16, 1.19; t = 2.57, p = 0.011; M3: linear estimate = 1.75, 95%CI: 0.73, 2.79; t = 3.36, p < 0.001). Relative reward was overall more susceptible to contagion versus absolute reward (M1: linear estimate = 1.40, 95%CI: 0.88, 1.92; t = 5.34, p<0.001; M3: linear estimate = 2.60, 95%CI: 1.57, 3.63; t = 4.98, p < 0.001). There was an interaction between group and belief type under M3 (M3: linear estimate = 2.13, 95%CI: 0.09, 4.18, t = 2.06, p=0.041) but not M1. There was a main effect of belief type on precision under M3 (linear estimate = 0.47, 95%CI: 0.07, 0.87, t = 2.34, p = 0.02) but not M1; relative reward preferences became more precise across the board. Derived model estimates of preference change between phase 1 and 3 strongly correlated between M1 and M3 along both belief types (see Table S2 and Fig S11). As a whole, humans are more susceptible to changing relative preferences more than absolute reward preferences, and this is disrupted in BPD.

Model Accuracy.

(A) We used random-effects hierarchical model fitting and comparison to jointly estimate group level and individual level parameters on simulated data (Piray et al., 2019). CON participants were best fit by M1, whereas BPD participants were best fit by M4 (B) Server matching between participant and partner in phase two was successful, with participants being approximately 50% different to their partners with respect to the choices each would have made on each trial in phase 2 (mean similarity=0.49, SD=0.12). Model accuracy across the task was very high (mean accuracy=0.8, SD=0.12). Model accuracy within each phase was very high (mean accuracy[phase1]=0.83, SD[phase1]=0.16; mean accuracy[phase2]=0.77, SD[phase2]=0.14; mean accuracy[phase3]=0.82, SD[phase3]=0.17). Loglikelihood values were also well above what would be expected had the model fitted the data by chance (median=-40.68, SD=22.7; chance value=-87.33). Choice probabilities generated by the model on each trial were also well above chance thresholds (median=0.91, SD=0.24; chance value=0.5). (C) The spearman association between the responsibility allocated for each participant during real and recovered model comparison was highly correlated on the diagonal. There was some correlation between M1-M2 but this was due to M2 being a nested model of M1, sharing similar free parameters; this was not worrying in light of excellent model identifiability overall in the synthetic comparison. Associations between real and recovered parameters from the dominant model within each BPD and CON participants was very high with few cross correlations on the off-diagonal. In both confusion and parameter recovery matrices, white spaces indicate insignificant associations at the p > 0.01 level. (D) (top panel) The relationship between uncertainty over the self and uncertainty over the other with respect to the change in the precision (left) and median-shift (right) in phase 3 relative reward preferences ( β^ppt ). CON participant self and other uncertainty is overlaid onto the plot to demonstrate the degree to which their beliefs should change in phase 3 according to the model. (bottom panel) Correlating the model-predicted median shift in beliefs and derived change in beliefs between phase 1 and 3 demonstrates a very strong association (r = 0.88, p < 0.001). For the purposes of visualisation we cap real and simulated values <15 for compactness, although the true correlation reported is irrelevant to this visual constraint. (E) (left panel) We overlay model-predicted (solid line) and real observed (dashed line) trial-by-trial probabilities extracted from a linear model for a correct prediction by participants. For raw trial by trial updating see Supplementary Figure 5. Both closely match. (middle panel) There was no significant difference (ns) for BPD and CON participants with respect to their total correct answers over phase 2. (right panel) Model-predicted and real observations in phase 2 total scores were highly correlated in both groups (CON r=0.84, p<0.001; BPD r=0.89, p<0.001).

Added in the revised version: Exploratory Psychometric and Intentional Attribution Analysis

Removed in the revised version: Exploratory Parameter Associations

Earlier version: We explored whether social contagion may be restricted as a result of childhood trauma, paranoia, and less effective trait mentalizing. We collected psychometric data from participants prior to entering the task and asked participants to attribute explicit intentions to their partner after phase 2. For raw parameter correlations with all psychometric measures see Figure S7. All analyses were corrected for False Discovery Rate (FDR; p[fdr]).

Revised version: Childhood trauma, persecution, and poor mentalising in BPD are all predicted to disrupt the integration of information from others (Fonagy & Luyten, 2009). Therefore we explored whether social contagion may be restricted as a result of childhood trauma, paranoia, and less effective trait mentalizing. We collected psychometric data from participants prior to entering the task and asked participants to attribute explicit intentions to their partner after phase 2. All analyses were corrected for False Discovery Rate (FDR; p[fdr]) and we provide correction for group status (Table S4).

Earlier version: We explored psychometric associations with social contagion under the assumption of M3 for all participants (where everyone is able to be influenced by their partner). We conducted partial correlation analyses to estimate relationships conditional on all other associations and retained all that survived bootstrapping (1000 reps), permutation testing (1000 reps), and subsequent FDR correction. Persecution and CTQ scores were both moderately associated with MZQ scores (RGPTSB r = 0.41, 95%CI: 0.23, 0.60, p=0.004, p[fdr]=0.043; CTQ r = 0.354 95%CI: 0.13, 0.56, p=0.019, p[fdr]=0.02). MZQ scores were in turn moderately and negatively associated with shifts in prosocial-competitive preferences (Δ βpptm ) between phase 1 and 3 (r = −0.26, 95%CI: −0.46, −0.06, p=0.026, p[fdr]=0.043). CTQ scores were also directly and negatively associated with shifts in individualistic preferences (Δ αpptm ; r = −0.24, 95%CI: −0.44, − 0.13, p=0.052, p[fdr]=0.065). This provides some preliminary evidence that trauma impacts beliefs about individualism directly, whereas trauma and persecutory beliefs impact beliefs about prosociality through impaired trait mentalising (Figure 4A). Social contagion under M3 was highly correlated with contagion under M1 (Fig S11).

Revised version: We assessed conditional psychometric associations with social contagion under the assumption of M3 for all participants. We conducted partial correlation analyses to estimate relationships conditional on all other associations and retained all that survived bootstrapping (5000 reps), permutation testing (5000 reps), and subsequent FDR correction. When not controlled for group status, RGPTSB and CTQ scores were both moderately associated with MZQ scores (RGPTSB r = 0.41, 95%CI: 0.23, 0.60, p[fdr]=0.043; CTQ r = 0.354 95%CI: 0.13, 0.56, p[fdr]=0.02). This was not affected by group correction. CTQ scores were moderately and negatively associated with shifts in individualistic reward preferences (Δ αpptm ; r = -0.25, 95%CI: -0.46, -0.04, p[fdr]=0.03). This was not affected by group correction. MZQ scores were in turn moderately and negatively associated with shifts in prosocial-competitive preferences (Δ βpptm ) between phase 1 and 3 (r = -0.26, 95%CI: -0.46, -0.06, p[fdr]=0.03). This was diminished when controlled for group status (r = 0.13, 95%CI: -0.34, 0.08, p[fdr]=0.20). Together this provides some evidence that self-reported trauma and self-reported mentalising influence social contagion (Fig S11). Social contagion under M3 was highly correlated with contagion under M1 demonstrating parsimony of outcomes across models (Fig S12).

Earlier version: We also tested parameter influences on explicit intentional attributions in Phase 2. Attributions included the degree to which they believed their partner was motived by harmful intent (HI) and self-interest (SI). According with prior work (Barnby et al., 2022), greater participant-partner disparity at the start of phase 2 was distinctly associated with HI and SI (Figure 4B), although associations did not survive FDR correction. Greater disparity of absolute preferences before learning was associated with reduced attributions of SI (ρ[| αparm − αpptm |] = −0.23, p=0.02, p[fdr]=0.08), and greater disparity of relative preferences before learning exaggerated attributions of HI (ρ[| βparm − βpptm |] = 0.21, p=0.04, p[fdr]=0.08). This is likely due to partners being significantly less individualistic and prosocial on average compared to participants (΄μ[α] = −5.50, 95%HDI: −7.60, −3.60; Δμ[β] = 12, 95%HDI: 9.70, 14.00); partners are correctly recognised as less selfish and more competitive.

Revised version: Prior work has predicted that partner-participant preference disparity influences mental state attributions (Barnby et al., 2022; Panizza et al., 2021). We tested parameter influences on explicit intentional attributions in Phase 2 while controlling for group status. Attributions included the degree to which they believed their partner was motived by harmful intent (HI) and self-interest (SI). According with prior work (Barnby et al., 2022), greater disparity of absolute preferences before learning was associated with reduced attributions of SI (ρ[| αparm − αpptm |] = −0.23, p[fdr]=0.08), and greater disparity of relative preferences before learning exaggerated attributions of HI (ρ[| βparm − βpptm |] = 0.21, p[fdr]=0.08), but did not survive correction (Figure S4B). This is likely due to partners being significantly less individualistic and prosocial on average compared to participants (΄μ[α] = −5.50, 95%HDI: −7.60, −3.60; Δμ[β] = 12, 95%HDI: 9.70, 14.00); partners are recognised recognised as less selfish and more competitive.

Earlier version: Greater prior uncertainty (before interaction) over a partner’s relative preferences was associated with increased HI (ρ[ βparref ] = 0.25, p=0.01, p[fdr]=0.04) but not SI, according with prior work (Barnby et al., 2022). There was no uncorrected association between prior uncertainty over absolute preferences with either attribution. This suggests that expectations of greater difference, irrespective of one’s true difference between self-other, may exaggerate beliefs about the intentional harm of others.

Revised version: Uncertainty in one’s expectations is associated with subsequent mental state attributions of intentional harm (Barnby et al., 2022). Greater prior uncertainty (before interaction) over a partner’s relative preferences was associated with increased HI (ρ[ βparref ] = 0.25, p[fdr]=0.04) but not SI, corrected for group status. There was no uncorrected association between prior uncertainty over absolute preferences with either attribution. Thus, expectations of greater difference between self-other, may exaggerate beliefs about the intentional harm of others.

Removed in the revised version: Psychometric correlations.

Removed in the revised version: (A) We conducted partial correlations between MZQ, CTQ, RGPTSB and changes in αpptm ; βpptm between phase 1 and 3 across the entire population under M3 to allow parameter comparison on the same scale. We performed bootstrapped resampling (1000 reps) and kept the edges that survived for visualisation. We then performed permutation analysis (1000 reps) to control for multiple comparisons within the partial network. Blue edges = positive partial correlation; Red edges = negative partial correlation. Note: the edge between the CTQ and Δ αpptm missed conventional significance after correction (p=0.065) (B). Uncorrected ranked spearman correlations between belief flexibility (| αparref ; βparref ) and absolute partner-participant dissimilarity (| αparm − αpptm |; | βparm − βpptm |) − with respect to server-derived parameters – in phase 2, controlling for group status. Only flexibility over relative reward preferences in phase 2 was associated with harmful intent attributions. Increased absolute participant-partner dissimilarity was associated with lower self-interest attributions, and increased relative participant-partner dissimilarity was associated with high harmful intent attributions.

Discussion

Earlier version: We built and tested a theory of interpersonal generalisation in a population of matched participants with (BPD) and without (CON) a diagnosis of borderline personality disorder using the Intentions Game, a three-phase social value orientation task. Both groups demonstrated equivalent behavioural accuracy but employed different strategies. CON participants used a process of self-other generalization to predict and align with their partners, while BPD participants maintained distinct representations of self and other, particularly over joint reward outcomes. As a whole, all participants were more sensitive to updates about joint versus absolute outcomes, with BPD participants particularly concerned with how outcomes relatively effected self and other. Our exploratory findings also indicate that retrospectively reported childhood trauma and persecutory beliefs were linked to reduced trait mentalising, which was subsequently associated with diminished self-change. Collectively, our results integrate prior findings in BPD and provide a formal account of social information generalisation in humans, alongside a concise social paradigm to test these processes.

Revised version: We built and tested a theory of interpersonal generalisation in a population of matched participants with (BPD) and without (CON) a diagnosis of borderline personality disorder using the Intentions Game, a three-phase social value orientation task. We compared four hypotheses, instantiated in formal computational models, to determine whether those with a diagnosis of BPD displayed disrupted self-insertion and social contagion. Both groups demonstrated equivalent behavioural accuracy but employed different strategies. CON participants used a process of self-other generalization to predict and align with their partners, while BPD participants maintained distinct representations of self and other, particularly over joint reward outcomes. As a whole, all participants were more sensitive to updates about joint versus absolute outcomes, with BPD participants particularly concerned with how outcomes relatively affected self and other. Our exploratory findings also indicate that retrospectively reported childhood trauma and persecutory beliefs were linked to reduced trait mentalising, which was subsequently associated with diminished shifts in participants’ relative reward preferences. Collectively, our results integrate prior findings in BPD and provide a formal account of social information generalisation in humans, alongside a concise social paradigm to test these processes.

The data replicate models of social generalisation that have focused on individual processes of self-insertion and contagion, extending these theories by demonstrating both processes in conjunction. Models of self-insertion directly map participant preferences onto prior beliefs about others, which has been used to explain increased reaction times in observational learning of others’ snack food preferences (Tarantola et al., 2017), as well as improved predictive accuracy when matched with individuals of similar social values (Barnby et al., 2022). Both findings are replicated in this study. Although we did not explicitly model reaction times, we observed an interaction between reaction time reductions over time and interpersonal similarity at baseline. In tandem, computational models of social contagion have focused on intertemporal discounting (Moutoussis et al., 2016) – with behavioural studies also focusing on effort-based reward (Devaine & Daunizeu, 2017) and moral preferences (Yu et al., 2021) - and explain shifts in self-preferences as a function of uncertainty regarding self and others. In both the dominant (M1) and sub-dominant (M3) models that best explained data in healthy participants, shifts in self-beliefs were also influenced by representational uncertainty of self and other: greater self-uncertainty and reduced other uncertainty led to larger shifts in social preferences.

Earlier version: The data also align with prior research on social impression formation, which suggests that humans form rapid evaluations of others that are refined over time (Bone et al., 2021; Moutoussis et al., 2023). This initial ‘heating’ and subsequent ‘cooling’ of beliefs corresponds to the computational complexity employed: model-based strategies are typically used early in interactions, transitioning to simpler, model-free computations once a partner’s behaviour becomes predictable (Gęsiarz & Crockett, 2015; Guennouni & Speekenbrink, 2022). Our findings support this framework, demonstrating initial variability early in interactions followed by steady updating.

Revised version: The data also align with prior research on social impression formation, which suggests that humans form rapid evaluations of others that are refined over time (Bone et al., 2021; Moutoussis et al., 2024). This initial ‘heating’ and subsequent ‘cooling’ of beliefs corresponds to the computational complexity employed: model-based strategies are typically used early in interactions, transitioning to simpler, model-free computations once a partner’s behaviour becomes predictable (Gęsiarz & Crockett, 2015; Guennouni & Speekenbrink, 2022). Our findings support this framework, demonstrating initial variability early in interactions followed by steady updating.

Disruptions in self-to-other generalization provide an explanation for previous computational findings related to task-based mentalizing in BPD. Studies tracking observational mentalizing reveal that individuals with BPD, compared to those without, place greater emphasis on social over internal reward cues when learning (Henco et al., 2020; Fineberg et al., 2018). Those with BPD have been shown to exhibit reduced belief adaptation (Siegel et al., 2020) along with ‘splitting’ of latent social representations (Story et al., 2024a). BPD is also shown to be associated with overgeneralisation in self-to-other belief updates about individual outcomes when using a one-sided reward structure (where participant responses had no bearing on outcomes for the partner; Story et al., 2024b). Our analyses show that those with BPD are equal to controls in their generalisation of absolute reward (outcomes that only affect one player) but disintegrate beliefs about relative reward (outcomes that affect both players) through adoption of a new, neutral belief. We interpret this together in two ways: 1. There is a strong concern about social relativity when those with BPD form beliefs about others, 2. The absence of self-insertion when predicting relative outcomes may predispose to brittle or ‘split’ beliefs. In other words, those with BPD assume ambiguity about the social relativity preferences of another (i.e. how prosocial or punitive) and are quicker to settle on an explanation to resolve this. Although self-insertion may be counter-intuitive to rational belief formation, it has important implications for sustaining adaptive, trusting social bonds via information moderation.

Those with a diagnosis of BPD also show reduced permeability in other-to-self generalising. While prior research has predominantly focused on how those with BPD use information to form impressions, it has not typically been examined whether these impressions affect the self. In interactive trust paradigms, neural responses to monetary offers from others to the self were substantially blunted in individuals with BPD compared to those without (King-Casas et al., 2008). Similarly, in non-social reward tasks, those with BPD show reduced neural feedback-related negativity amplitudes, which obstructs feedback-related self-change (Stewart et al., 2019; Vega et al., 2013). Our results suggest a mechanistic basis for social contagion, indicating that self-rigidity prevents observed social behaviours from generalizing to the self, potentially exacerbated by childhood trauma, paranoia, and impaired mentalizing capabilities. Resistance to social influence may serve as a protective response but can also contribute to the pervasive loneliness experienced by individuals with BPD, even in the absence of social isolation (Liebke et al., 2017).

Earlier version: Notably, despite differing strategies, those with BPD achieved similar accuracy to CON participant. While all participants were more concerned with relative vs. absolute reward, those with BPD changed their strategy contingent on this focus. Practically this difference in BPD is captured either through disintegrated priors with a new median (M4) or very noisy, but integrated priors over partners (M1) if we assume M1 can account for the full population. In either case, the algorithm underlying the computational goal for BPD participants is far higher in entropy and emphasises a less stable or reliable process of inference. It is important to assess this mechanism alongside momentary assessments of mood to understand whether more entropic learning processes contribute to distressing mood fluctuation.

Revised version: Notably, despite differing strategies, those with BPD achieved similar accuracy to CON participants. While all participants were more concerned with relative vs. absolute reward, those with BPD changed their strategy contingent on this focus. Practically this difference in BPD is captured either through disintegrated priors with a new median (M4) or very noisy, but integrated priors over partners (M1) if we assume M1 can account for the full population. In either case, the algorithm underlying the computational goal for BPD participants is far higher in uncertainty, whether through a neutral central tendency (M4) or large variance (M1) prior over relative reward in phase 2, and emphasises a less stable or reliable expectation about others. It is important to assess this mechanism alongside momentary assessments of mood to understand whether more entropic learning processes contribute to distressing mood fluctuation.

Clinical implications of our work underscore the importance of consistency and stability in clinical support for individuals with a diagnosis of BPD. Encouragingly, we found that those with BPD were not entirely impermeable to observed behaviour, suggesting that consistent external models of trust could be internalized over time. Restoring a stable sense of self through social learning and effective mentalizing (Nolte et al., 2023), along with a consistent focus on differentiating self from other (de Meulemeester et al., 2021), are central to mentalization-based therapies (Bateman & Fonagy, 2010; Smits et al., 2024) and other evidence-based treatments for BPD. We hope that our paradigm and model can offer insights into the effectiveness of these and other therapies in driving mechanistic psychological change. A key task for future work will be to assess whether generalisation principles may be restored in within-individuals with a diagnosis of BPD following intervention.

More broadly, our model bridges formal theories of associative learning and social cognition. Reinforcement learning approaches have effectively organized theories around uncertainty navigation in non-social contexts (Piray et al., 2021; Zika, 2023). However, humans do not function in isolation. Bayesian models of internal and external social beliefs are better suited to capture the dynamic nature of time, context, and uncertainty during interactions (FeldmanHall & Nassar, 2021; Velez & Gweon, 2021), where joint reward rather than individual reward may be particularly salient (Barnby et al., 2023). Our paradigm is concise, visually engaging, includes straightforward rules and instructions, and allows for tight experimental control over partner similarity. Our model and paradigm effectively capture core social psychological principles grounded in general computational approaches to learning and uncertainty, elucidating key aspects of human social interaction and exchange.

Earlier version: We note some limitations to our study. Primarily, we focused on the ability of individuals to integrate their self-concept into beliefs about others. It is also possible that humans possess strong, salient representations of others (or groups of others) that serve as dominant templates for learning. This may be particularly relevant for individuals with BPD, who will often have interpersonal experiences of abuse, neglect, or other forms of distress. The use of a salient, negative other-prior as a basis for learning was not measured in this study, but it may explain the ambivalent prior observed in phase 2, where a mixture of self and notional other influences belief formation, leading to rigid belief updating. Individuals with BPD may integrate priors from different sources as a mixture. We can simulate this by modelling a causal framework that incorporates priors based on both self and a strong memory impression of a notional other (Figure S3). However, a strength of our data is that we observed impression formation independent of valence—impressions were formed regardless of whether a partner was more or less prosocial or selfish than the participant (Figure S4). This supports our hypothesis that a vulnerable self-model and lack of self-insertion contribute to the formation of overly precise beliefs during learning as a means of rapidly reducing uncertainty. Even if a mixture model better explains the ambivalent prior in phase 2, it would still support a general hypothesis about the fractured concept of self and other in BPD.

Revised version: We note some limitations to our study. Primarily, we focused on the ability of individuals to integrate their self-concept into beliefs about others. It is also possible that humans possess strong, salient representations of others (or groups of others) that serve as dominant templates for learning. This may be particularly relevant for individuals with BPD, who will often have interpersonal experiences of abuse, neglect, or other forms of distress. The use of a salient, negative other-prior as a basis for learning was not measured in this study, but it may explain the ambivalent prior observed in phase 2, where a mixture of self and notional other influences belief formation, leading to rigid belief updating. Individuals with BPD may integrate priors from different sources as a mixture. We can simulate this by modelling a framework that incorporates priors based on both self and a strong memory impression of a notional other (Figure S3). However, a strength of our data is that we observed impression formation independent of valence—impressions were formed regardless of whether a partner was more or less prosocial or selfish than the participant (Figure S4). This supports our hypothesis that a vulnerable self-model and lack of self-insertion contribute to the formation of overly precise beliefs during learning as a means of rapidly reducing uncertainty. Even if a mixture model better explains the ambivalent prior in phase 2, it would still support a general hypothesis about the fractured concept of self and other in BPD.

Another strength of our work is demonstrating processes of self-insertion and contagion under minimal interaction conditions: simple observation alone was sufficient to elicit both processes. However, this is also a limitation. While we predict that these processes will apply in more naturalistic settings, this has yet to be tested, and it remains unclear whether these effects will persist in richer conditions, particularly when higher affective arousal and challenges to mentalising are present. Lastly, the action space and parameters governing choice in our study were quite simple—two actions influenced by two parameters. This was a deliberate computational choice to avoid overly complex action spaces that may be difficult to fit to real human data, and which might fail to capture how these mechanisms operate in the context of increasing action and model complexity. As a whole, our findings open new possibilities for testing how social uncertainty across the lifespan (e.g. in adolescence; Sebastian et al., 2008), and in the context of ill-health, may explain the formation and maintenance of healthy social bonds as well as their disruption.

Earlier version: Finally, a limitation may be that behaviour in tasks based on economic preferences may not have clinical validity. This issue is central to the field of computational psychiatry, much of which is based on generalising from tasks like that within this paper and discussing correlations with psychometric measures. Extrapolating economic tasks into the real world has been the topic of discussion for the many reviews on computational psychiatry (e.g. Montague et al., 2012; Hitchcock et al., 2022; Huys et al., 2016). We note a strength of this work is the use of model comparison to understand causal algorithmic differences between those with BPD and matched healthy controls. Nevertheless, we wish to further pursue how latent characteristics captured in our models may directly relate to real-world affective change.

Revised version: Finally, a limitation may be that behaviour in tasks based on economic preferences may not have clinical validity. This issue is central to the field of computational psychiatry, much of which is based on generalising from tasks like that within this paper and discussing correlations with psychometric measures. Extrapolating economic tasks into the real world has been the topic of discussion for the many reviews on computational psychiatry (e.g. Montague et al., 2012; Hitchcock et al., 2022; Huys et al., 2016). We note a strength of this work is the use of model comparison to understand algorithmic differences between those with BPD and matched healthy controls. Nevertheless, we wish to further pursue how latent characteristics captured in our models may directly relate to real-world affective change.

Materials and Methods

Participants

Earlier version: We used a case-control, between-subjects design with 103 participants: a control group from the general population (N = 53) and a clinical group diagnosed with BPD (N = 50). Both groups were recruited for a larger study investigating social exchanges in BPD and Anti-Social Personality Disorder (approved by the Research Ethics Committee for Wales, 12/WA/0283). The control and clinical groups were matched on age, sex, years in education, and the English Indices of Deprivation based on the 2019 census (IoD 2019; Ministry of Housing, Communities & Local Government, 2019). Participants received £70 compensation for completing questionnaires and online tasks which included the Intentions Game. They also received a performance bonus if they were entered into the lottery for surpassing 1000 points over the course of the game.

Revised version: We used a case-control, between-subjects design with 103 participants: a control group from the general population (N = 53) and a clinical group diagnosed with BPD (N = 50). Both groups were recruited for a larger study investigating social exchanges in BPD and Anti-Social Personality Disorder (approved by the Research Ethics Committee for Wales, 12/WA/0283). The control and clinical groups were matched on age, sex, years in education, and the English Indices of Deprivation based on the 2019 census (IoD2019; Ministry of Housing, Communities & Local Government, 2019). Participants received £70 compensation for completing questionnaires and online tasks which included the Intentions Game. They also received a performance bonus if they were entered into the lottery for surpassing 1000 points over the course of the game.

Participants for the control group were recruited through an advertisement on the Call For Participants website (https://www.callforparticipants.com), local community services and adult schools. Inclusion criteria required control participants to have no pre-existing or current diagnoses of mental health disorders, neurological disorders, or traumatic brain injuries. Additionally, control participants must not have been currently in therapy or taking medication for any psychiatric disorders.

The majority of BPD participants were recruited through referrals by psychiatrists, psychotherapists, and trainee clinical psychologists within personality disorder services across 9 NHS Foundation Trusts in the London, and 3 NHS Foundation Trusts across England (Devon, Merseyside, Cambridgeshire). Four BPD participants were also recruited by self-referral through the UCLH website, where the study was advertised. To be included in the study, all participants needed to have, or meet criteria for, a primary diagnosis of BPD (or emotionally-unstable personality disorder or complex emotional needs) based on a professional clinical assessment conducted by the referring NHS trust (for self-referrals, the presence of a recent diagnosis was ascertained through thorough discussion with the participant, whereby two of the four also provided clinical notes). The patient participants also had to be under the care of the referring trust or have a general practitioner whose details they were willing to provide. Individuals with psychotic or mood disorders, recent acute psychotic episodes, severe learning disability, or current or past neurological disorders were not eligible for participation and were therefore not referred by the clinical trusts.

Psychometric Measures

Green et al. Paranoid Thought Scale (GPTS)

The GPTS assesses paranoid thoughts, including ideas of social reference (scale A) and persecution (scale B), in both general and clinical populations (Green et al., 2008). Each item is scored from 0 (not at all) to 5 (totally) concerning endorsement of each item. We retained items from the GPTS that were consistent with the revised version outlined in Freeman et al., 2021 (Revised GPTS; R-GPTS). The R-GPTS has demonstrated excellent psychometric properties (Freeman et al., 2021), making it a reliable and valid tool for assessing trait paranoid thoughts in non-clinical and clinical populations.

Childhood Trauma Questionnaire (CTQ)

The Childhood Trauma Questionnaire is used to screen for maltreatment history (Bernstein et al., 2003). Each item is scored from 1 (never true) to 5 (very often true). The CTQ has showed good internal consistency reliability across the five scales (Sacchi et al., 2018) and good construct validity based on significant associations with stress responsivity (McMahon et al., 2022), and dissociation (Nobakht et al., 2021).

Certainty About Mental States Questionnaire (CAMSQ)

The CAMSQ assesses one’s certainty in classifying the mental states of oneself and others at an abstract level (Müller et al., 2023), e.g. ‘I know what other people think of me’ and ‘I know my feelings’. Each subscale is scored from 1 (never) to 7 (always). In US and German samples, the CAMSQ showed high internal consistency for Self-Certainty (ω = .90/.88) and Other-Certainty (ω = .91/.89) subscales, and high two-week test-retest reliability for Self-Certainty (r = .85), Other-Certainty (r = .78), and Other-Self-Discrepancy (r = .82) scores (Müller et al., 2023).

Mentalisation Questionnaire (MZQ)

The MZQ is a 15-item questionnaire assessing an individual’s trait mentalizing, i.e., one’s ability to understand and interpret their own and others’ mental states (Hausberg et al., 2012). The MZQ demonstrated good internal consistency (α = .81) and test-retest reliability (r = .76), and was sensitive to change over a 6-month follow-up period and showed good criterion-related validity, distinguishing individuals with BPD from those without BPD (Hausberg et al., 2012). A higher score reflects worse trait mentalizing.

Epistemic Trust, Mistrust and Credulity Questionnaire

The ETMCQ is a 15-item measure calibrated to assess trust (e.g. ‘I usually ask people for advice when I have a personal problem), mistrust (e.g. ‘I’d prefer to find things out for myself on the internet rather than asking people for information), and credulity (e.g. ‘I am often considered naïve because I believe almost anything that people tell me’; Campbell et al., 2021). Each item is scored from 1(Strongly Disagree) to 7(Strongly Agree).

Paradigm, procedure and server architecture

The Intentions Game is a repeated social-value orientation paradigm with three phases.

In Phase 1 of the Intentions Game, participants take on the role of the decider with an anonymous partner over 36 trials. In each trial, participants choose between two options to distribute points between themselves and their partners. Participants make 12 choices each between prosocial and competitive (e.g. Option 1=[10,10], Option 2 = [10,5]) individualistic and competitive (e.g. Option 1=[10,5], Option 2=[8,1]), and prosocial and individualistic options (e.g. Option 1=[5,5], Option 2=[10,5]). Phase 1 choices allowed experimenters to classify participants’ social preferences as prosocial (preferring equal outcomes), individualistic (maximising own payoff), or competitive (maximising relative payoff difference at the cost of lower self-gain).

We included a task environment that balanced each type of choice pair (see Supplementary Table 1).

In phase 2 of the game, participants were matched with a new anonymous partner and played the role of the recipient over 54 trials. In this phase, the participants predicted which of the two options their partner would choose on each trial. Trial numerical values for self and other were identical to Phase 1. Partners’ decisions were determined via a dynamic algorithm (Burgess et al., 2023) to ensure partners were approximately ~50% different from the participants’ based on participants’ choices in phase 1. To surmise this architecture, we implemented a version of the client-server paradigm hosted on an Amazon Web Service (AWS) LightSail server, where the webbased behavioural task (implemented with JavaScript in Gorilla.sc) acted as the client and exchanged information with a remote AWS server. The server received all anonymised behavioural data following phase 1. The Application Programming Interface (API) to interact with the server used a customizable R script (v4.3) to process the received data from the participant, and additional R scripts were used to process and generate output for the participant. A function within the backend scripts first used Bayesian inference to approximate a participant’s parameters for phase 1. It then simulated what choices the participant would have made in phase 2 had the participant been in the role of the partner. The algorithm then sought to find parameters that would be at least 50% dissimilar from participant parameters with respect to the generated choices of those parameters. This allowed the task behaviour of phase 2 to be dynamically updated in response to participant choices in phase 1. This facilitated tight control over the state of the task and enabled advanced computations to be performed on participant data beyond the capabilities of a web browser.

Participants were incentivised in phase 2 to predict accurately, as accurate predictions would contribute to their total point scores (total correct answers were multiplied by 10 and added to their points) and determined their entry into the lottery to win an extra £20 Amazon voucher. After participants had made their predictions, they were given feedback informed on whether their predictions were accurate.

At the end of phase 2, participants were asked to rate (1) the extent to which they thought their partner was driven by the desire to earn points in this task overall (self-interest) and (2) the extent to which they thought their partner was driven by the desire to reduce the participant’s points in this task overall (attribution of harmful intent). The answers were presented using two separate sliders from 0 to 100; the sliders were initialised to be invisible until the participants made the first click.

Phase 3 was identical to phase 1 except that participants were matched with a new anonymous partner. Participants would take on the decider role similar to phase 1 which allowed experimenters to estimate whether the observation of their partner in phase 2 had an influence on participants in phase 3.

Behavioural Analysis

Earlier version: All analysis was conducted in R (v. 4.3.3) on a macbook pro (M2 Max; OS=Ventura13.5). All individual numeric values extraneous of statistical tests are reported with their mean and standard deviation (mean=KCWM8V2, SD=YY). All statistical tests where dependent variables mapped one value to one participant (e.g. trait psychometric scores) were conducted as linear models, with the regression coefficient, 95% confidence interval (95%CI), t-value and p-value reported like so (linear estimate=KCWM8V2, 95%CI:AA,BB; t=CC, p=DD). When dependent variables mapped multiple values to each participant (e.g. trial-by-trial accuracy or reaction time) random-effects linear modelling was used. All correlations used Pearson estimates (r) unless distributions were non-normal, in which case Spearman-ranked correlations (r) were performed.

Revised version: All analysis was conducted in R (v. 4.3.3) on a macbook pro (M2 Max; OS=Ventura13.5). All individual numeric values extraneous of statistical tests are reported with their mean and standard deviation (mean=XX, SD=YY). All statistical tests where dependent variables mapped one value to one participant (e.g. trait psychometric scores) were conducted as linear models, with the regression coefficient, 95% confidence interval (95%CI), t-value and p-value reported like so (linear estimate=XX, 95%CI:AA,BB; t=CC, p=DD). When dependent variables mapped multiple values to each participant (e.g. trial-by-trial accuracy or reaction time) random-effects linear modelling was used. All correlations used Pearson estimates (r) unless distributions were non-normal, in which case Spearman-ranked correlations (r) were performed.

Model space and Computational Analysis

Earlier version: We apply four computational hypotheses (M1-M4) which could explain the data collected from the Intentions Game (Figure 1), centred around formal principles of self-insertion and social contagion. Self-insertion states that a self inserts their own preferences into their beliefs about others (Anderson & Chen, 2002; Kreuger & Clement, 1994); Social Contagion states that a self’s preferences will change when exposed to the preferences of an other (Frith & Frith, 2012). In each case, cognitive representations of self and other are allowed to intermingle to form a new hybrid of the two for the purposes of computational efficiency and/or social bonding.

Revised version: We apply four computational hypotheses (M1-M4) which could explain the data collected from the Intentions Game (Figure 1), centred around formal principles of self-insertion and social contagion. Self-insertion states that a self inserts their own preferences into their beliefs about others (Anderson & Chen, 2002; Kreuger & Clement, 1994); Social Contagion states that a self’s preferences will change when exposed to the preferences of an other (Deutsch & Gerard, 1955). In each case, cognitive representations of self and other are allowed to intermingle to form a new hybrid of the two for the purposes of computational efficiency and/or social bonding.

We note some important assumptions in our notation going forward. In dyadic social interaction, both parties are trying to estimate and predict the true state (θ) of the self (θS) and the other (θO). However, this estimation is inherently imperfect. Theories of social inference need to consider three sources of noisy estimation of this quantity: the self’s (s) metacognitive model of their own state, θ¯s , their partner’s (O) state, θ¯o , and finally the experimenter’s approximation of both quantities, θ¯s,o , (Barnby et al., 2024). In this work we consider the experimenter’s approximation of the self’s state θ^s (phase1), the self’s approximation of their other θ¯o (phase 2), and how exposure to a partner may influence θ^s (phase 3). We term the self the participant (ppt) and the other the partner (par) and assume θppt = θ^s in phase 1, θpar = θ¯o in phase 2, and θ′ppt are the shifted participant preferences following exposure to the partner.

All models assumed a constricted Fehr-Schmidt utility function was used by participants and partners to calculate the utility of two options (Uα,β={ Uα,β1,Uα,β2 }) in each trial within the task.

In phase 1, participants made binary choices ct, t = {1… T} about whether option 1 or option 2 should be chosen given the returns for each option pair, Rt = {Rt;1; Rt;2} = { Rpptt;1,Rpart;1Rpptt;2,Rpart;2 } .

Earlier version: Here, αppt describes the weight a participant places on their own payoff (in one reduced model we set αppt = 1), and βppt, the weight a participant places on their payoff relative to the payoff of their partner. Large positive or negative values of βppt indicate respectively that participants like or dislike earning more than their partner. We can therefore describe these terms α and β as reflecting preferences for absolute and relative payoffs, respectively. For efficiency we discretised states of αppt from 030 (increments of 0.125) and βppt from −30 to 30 (increments of 0.25).

Revised version: Here, αppt describes the weight a participant places on their own payoff (in one reduced model we set αppt = 1), and βppt, the weight a participant places on their payoff relative to the payoff of their partner. Large positive or negative values of βppt indicate respectively that participants like or dislike earning more than their partner. We can therefore describe these terms α and β as reflecting preferences for absolute and relative payoffs, respectively. For efficiency we discretised states of αppt from 0-30 (increments of 0.125) and βppt from −30 to 30 (increments of 0.25).

Over this state space we can construct a belief that participants are estimated to hold which generate their choices, C. Herein, we refer to this belief as θppt, where θppt is a matrix over a fixed grid of αppt and βppt values. In the models, θppt is drawn from a normal distribution made from a central tendency, θpptm , and a standard deviation, θpptσ . The standard deviation around the central tendency allows for stochastic choice behaviour consistent with random utility models (Block, 1974; McFadden, 1974). We invert the model to estimate θppt based on a participant’s choices given their likelihood of choosing ct = 1 : p(θppt|C)~N(θppt;θpptm,θpptσ)p(ct=1|θppt,Rt)=∑θpptσ(ΔUα,β(Rt))⋅θpptLL=log[ p(ct=1|θppt,Rt) ] (2)

When θpptσ is larger, a participant’s choices in Phase 1 are estimated to be less deterministic and more stochastic – i.e. they are less sure about their preferences along each dimension. This consideration will become important for choices made in phase 3.

In phase 2, over 54 trials, we then model the participants binary predictions d¯t , t = {1… T} about whether option 1 or 2 would be chosen by their partner given the returns Rt = {Rt;1; Rt;2} = { Rpptt;1,Rpart;1,Rpptt;2,Rpart;2 } for each pair of options. They then were given feedback about the partner’s true decision which we note as df. We assumed the participant predict the partner in the same way they would themselves, ranging along two dimensions, αpar and βpar which was needed to be inferred through observation, using a likelihood for df of LL = log [p(dt = 1|αpar, βpar, Rt)] using the same formula as phase 1. We note the belief about αpar and βpar together as θpar, represented as a matrix over a fixed grid of αpar and βpar values.

The partner decisions, Dt = {d1, d2 …, dT} are then used to update the participants beliefs about the partner, written as p(θpar|Df), starting with prior p(θpar|D0). Both M1 and M2 assume participants use their own central tendency, θpptm , as a starting point for their prior beliefs about their partner as theoretically outlined as a self-insertion bias (Barnby et al., 2024) which draws from past computational work (Barnby et al., 2022; Tarantola et al., 2017). We also assumed participants used a new standard deviation θparref which allowed for participants to believe their partner may be different from them (belief flexibility). Therefore we have: p(θpar|D0)~N(θpar;θpptm,θparref) (3)

In models M3 and M4, we assume participant’s may instead use a new central tendency (rather than their own) as prior beliefs over their partner. This are free parameters to be approximated, α¯parm , β¯parm .

In all cases, we assume participants update their beliefs about their partner’s social preferences given their partner’s decisions D along trials 1-54 according to Bayes rule: θpart=p(dt|θpar;Rt)θpart−1∑θ′parp(dt|θ′par;Rt)θ′part−1 (4)

We can then marginalise over θpart to calculate the belief participants had over their participant’s social value preferences.

We assume that participants predict the partner’s decision in the next trial by calculating the probability determined by the utility differences ΔUα, β(Rt+1) as in phase 1, summed over the joint distribution of partner parameters, θpart : p(d¯t+1=1|Dt,Rt)=∑θparσ(ΔUα,β(Rt))⋅θpartp(d¯t+1=2|Dt,Rt)=1−p(d¯t+1=1|Dt,Rt) (5)

And then performed probability matching, so that: p(d¯t+1=1|Dt,Rt)=p(dt+1=1|Dt,Rt) (6)

In the third phase participants are once again asked to make choices for themselves and a new anonymous partner over 36 trials with an assumed identical utility function as in phase 1. In model M1 and M3 we assume participants use a combination of their own preferences and the posterior beliefs about their partner to form a new distribution to select between the two options available on each trial. This draws from the same formulation used previously (Moutoussis et al. 2016). In essence, we state that participants know their true preferences in phase 1 but are unsure about them. The inferred partner beliefs θpart provides information to the participant about some common preference distribution both share, which in turn informs the participant’s own choices c′t , t = {1… T} in the form of an adjusted belief along each dimension for phase 3, α′ppt and β¯ppt (eq. 7), using a log likelihood of LL = log[p( c′t = 1| α′ppt , β′ppt , Rt)]. We refer to α′ppt and β′ppt together as θ′ppt for convenience, where θ′ppt is a matrix over a grid of fixed values of α′ppt and β′ppt . To note: models M2 and M4 do not assume participants undergo this change, and instead use their original phase 1 beliefs to make choices LL = log[p( c′t = 1|αppt, βppt, Rt)].

Where θparσ and θparm are the standard deviation and central tendency of the final posterior inference about the partner, θpart;54 .

All computational models were fitted using a Hierarchical Bayesian Inference (HBI) algorithm which allows hierarchical parameter estimation while assuming random effects for group and individual model responsibility (Piray et al., 2019). During fitting we added a small noise floor to distributions (2.22e−16) before normalisation for numerical stability. Parameters were estimated using the HBI in untransformed space drawing from broad priors (μM=0, σ2M = 6.5; where M={M1, M2, M3, M4}). This process was run independently for each group. Parameters were transformed into model-relevant space for analysis. All models and hierarchical fitting were implemented in Matlab (Version R2022B). All other analyses were conducted in R (version 4.3.3; arm64 build) running on Mac OS (Ventura 13.0). We extracted individual and group level responsibility, as well as the protected exceedance probability to assess model dominance per group.

To conduct model recovery we simulated synthetic participants (CON=53; BPD=50) using their fitted parameters from the dominant model of the group (CON=M1; BPD=M4). We then performed model fitting with an identical procedure to the real behavioural data. We tested associations between model responsibility and individual parameters for the real and recovered models, as well as the association between choices and predictions made by the model from simulation and the choices and predictions made by participants in each trial.

Differences between groups for individual-level parameters were estimated using hierarchical Bayesian t-tests (Bååth, 2014) and hierarchical general linear models in rStanArm. Differences in mean between groups (Δμ) are additionally reported with their corresponding posterior 95% High Density Interval (95%HDI). Belief updates were calculated as the Kullback-Leibler Divergence between probabilities (P) from trial t-1 to t, marginalised along all possible states, S={s1, s2, … , sn}: DKL(Pf||Pt-1) = ∑sSPt(s)logPt(s)Pt−1(s) .

Added in the revised version: Exploratory Network Analysis

Added in the revised version: To understand the individual differences of trait attributes (MZQ, RGPTSB, CTQ) with other-to-self information transfer (Δ αpptm ; Δ βpptm ) across the entire sample we performed a network analysis (Borsboom, 2021). Network analysis allows for conditional associations between variables to be estimated; each association is controlled for by all other associations in the network. It also allows for visual inspection of the conditional relationships to get an intuition for how variables are interrelated as a whole (see Fig S11). We implemented network analysis with the bootNet package in r using the ‘estimateNetwork’ function with partial correlations (Epskamp, Borsboom & Fried, 2018). To assess the stability of the partial correlations we further implemented bootstrap resampling with 5000 repetitions using the ‘bootnet’ function. We then additionally shuffled the data and refitted the network 5000 times to determine a ppermuted value; this indicates the probability that a conditional relationship in the original network was within the null distribution of each conditional relationship. We then performed False Discovery Rate correction on the resulting p-values. We additionally controlled for group status for all variables in a supplementary analysis (Table S4).